---
  title: "Test"
  description: "Test"
  v2: true

---
## Sample exercise

```yaml
type: NormalExercise
lang: python
xp: 100
skills: 2
key: 91cf6ff10e



```

Test test

`@instructions`
test test test

`@hint`
test

`@pre_exercise_code`
```{python}
print('test')
```
`@sample_code`
```{python}
print('test')
```
`@solution`
```{python}
print('test test)
```
`@sct`
```{python}
test
```




